# HW2: Khám phá dữ liệu, tiền xử lý, phân tích đơn giản
**Vì đây là bài tập về Pandas nên yêu cầu là không được dùng vòng lặp**


Họ tên: Lưu Trường Dương

MSSV: 19120489

---

## Cách làm bài và nộp bài (bạn cần đọc kỹ)

&#9889; Bạn lưu ý là mình sẽ dùng chương trình hỗ trợ chấm bài nên bạn cần phải tuân thủ chính xác qui định mà mình đặt ra, nếu không rõ thì hỏi, chứ không nên tự tiện làm theo ý của cá nhân.

**Cách làm bài**

Bạn sẽ làm trực tiếp trên file notebook này. Đầu tiên, bạn điền họ tên và MSSV vào phần đầu file ở bên trên. Trong file, bạn làm bài ở những chỗ có ghi là:
```python
# YOUR CODE HERE
raise NotImplementedError()
```
hoặc đối với những phần code không bắt buộc thì là:
```python
# YOUR CODE HERE (OPTION)
```
hoặc đối với markdown cell thì là:
```markdown
YOUR ANSWER HERE
```
Tất nhiên, khi làm thì bạn xóa dòng `raise NotImplementedError()` đi.
Đối những phần yêu cầu code thì thường ở ngay phía dưới sẽ có một (hoặc một số) cell chứa các bộ test để giúp bạn biết đã code đúng hay chưa; nếu chạy cell này không có lỗi gì thì có nghĩa là qua được các bộ test. Trong một số trường hợp, các bộ test có thể sẽ không đầy đủ; nghĩa là, nếu không qua được test thì là code sai, nhưng nếu qua được test thì chưa chắc đã đúng.

Trong khi làm bài, bạn có thể cho in ra màn hình, tạo thêm các cell để test. Nhưng khi nộp bài thì bạn xóa các cell mà bạn tự tạo, xóa hoặc comment các câu lệnh in ra màn hình. Bạn lưu ý <font color=red>không được tự tiện xóa các cell hay sửa code của Thầy</font> (trừ những chỗ được phép sửa như đã nói ở trên).

Trong khi làm bài, thường xuyên `Ctrl + S` để lưu lại bài làm của bạn, tránh mất mát thông tin.


*Nên nhớ mục tiêu chính ở đây là <font color=green>học, học một cách chân thật</font>. Bạn có thể thảo luận ý tưởng với bạn khác cũng như tham khảo các nguồn trên mạng, nhưng sau cùng <font color=green>code và bài làm phải là của bạn, dựa trên sự hiểu thật sự của bạn</font>. Khi tham khảo các nguồn trên mạng thì bạn cần ghi rõ nguồn trong bài làm. Bạn không được tham khảo bài làm của các bạn năm trước (vì nếu làm vậy thì bao giờ bạn mới có thể tự mình suy nghĩ để giải quyết vấn đề); sau khi kết thúc môn học, bạn cũng không được đưa bài làm cho các bạn khóa sau hoặc public bài làm trên Github (vì nếu làm vậy thì sẽ ảnh hưởng tới việc học của các bạn khóa sau). Nếu bạn có thể làm theo những gì mình nói thì điểm của bạn có thể sẽ không cao nhưng bạn sẽ có được những bước tiến thật sự. <font color=red>Trong trường hợp bạn vi phạm những điều mình nói ở trên thì sẽ bị 0 điểm cho toàn bộ môn học.</font>*

**Cách nộp bài**

Khi chấm bài, đầu tiên mình sẽ chọn `Kernel` - `Restart & Run All`, để restart và chạy tất cả các cell trong notebook của bạn; do đó, trước khi nộp bài, bạn nên chạy thử `Kernel` - `Restart & Run All` để đảm bảo mọi chuyện diễn ra đúng như mong đợi.

Sau đó, bạn tạo thư mục nộp bài theo cấu trúc sau:
- Thư mục `MSSV` (vd, nếu bạn có MSSV là 1234567 thì bạn đặt tên thư mục là `1234567`)
    - File `Lab02_KhamPhaPhanTichDuLieu.ipynb` (không cần nộp các file khác)

Cuối cùng, bạn nén thư mục `MSSV` này lại và nộp ở link trên moodle. Đuôi của file nén phải là .zip (chứ không được .rar hay gì khác).

<font color=red>Bạn lưu ý tuân thủ chính xác qui định nộp bài ở trên.</font>

---

## Import

```python
import matplotlib.pyplot as plt
import numpy as np
import pandas as pd 
# YOUR CODE HERE (OPTION) 
# Nếu cần các thư viện khác thì bạn có thể import ở đây
```

---

## Thu thập dữ liệu

Dữ liệu được sử dụng trong bài tập này là dữ liệu khảo sát các lập trình viên của trang StackOverflow. Mình download dữ liệu [ở đây](https://drive.google.com/file/d/1dfGerWeWkcyQ9GX9x20rdSGj7WtEpzBB/view) và có bỏ đi một số cột để đơn giản hóa. 

Theo mô tả trong file "README_2020.txt" của StackOverflow:
>The enclosed data set is the full, cleaned results of the 2020 Stack Overflow Developer Survey. Free response submissions and personally identifying information have been removed from the results to protect the privacy of respondents. There are three files besides this README:
>
>1. survey_results_public.csv - CSV file with main survey results, one respondent per row and one column per answer
>2. survey_results_schema.csv - CSV file with survey schema, i.e., the questions that correspond to each column name
>3. so_survey_2020.pdf - PDF file of survey instrument
>
>The survey was fielded from February 5 to February 28, 2020. The median time spent on the survey for qualified responses was 16.6 minutes.
>
>Respondents were recruited primarily through channels owned by Stack Overflow. The top 5 sources of respondents were onsite messaging, blog posts, email lists, Meta posts, banner ads, and social media posts. Since respondents were recruited in this way, highly engaged users on Stack Overflow were more likely to notice the links for the survey and click to begin it.

File "survey_results_public-short.csv" mà mình đính kèm là phiên bản đơn giản hóa của file "survey_results_public.csv" (từ 61 cột, mình bỏ xuống còn 29 cột). Đây là file dữ liệu chính mà bạn sẽ làm trong bài tập này. Ngoài ra, mình còn đính kèm 2 file phụ: (1) file "survey_results_schema-short.csv" là file cho biết ý nghĩa của các cột, và (2) file "so_survey_2020.pdf" là file khảo sát gốc của StackOverflow.

Để ý: 
- Dữ liệu này không đại diện được cho cộng đồng lập trình viên trên toàn thế giới, mà chỉ giới hạn trong tập những lập trình viên thực hiện khảo sát của StackOverflow. Những câu trả lời có được thông qua tập dữ liệu này cũng sẽ bị giới hạn trong phạm vi đó.
- Dữ liệu có đúng không? Về cơ bản là ta không biết được. Ở đây, mục đích chính là học qui trình Khoa Học Dữ Liệu và các câu lệnh của Pandas nên ta sẽ **giả định** phần lớn dữ liệu là đúng và tiếp tục làm.

Cũng theo file "README_2020.txt", dữ liệu này được StackOverflow public với license như sau:
>This database - The Public 2020 Stack Overflow Developer Survey Results - is made available under the Open Database License (ODbL): http://opendatacommons.org/licenses/odbl/1.0/. Any rights in individual contents of the database are licensed under the Database Contents License: http://opendatacommons.org/licenses/dbcl/1.0/
>
>TLDR: You are free to share, adapt, and create derivative works from The Public 2020 Stack Overflow Developer Survey Results as long as you attribute Stack Overflow, keep the database open (if you redistribute it), and continue to share-alike any adapted database under the ODbl.


---

## Khám phá dữ liệu

### Đọc dữ liệu từ file (0.25đ)

Đầu tiên, bạn viết code để đọc dữ liệu từ file "survey_results_public-short.csv" và lưu kết quả vào DataFrame `survey_df`; ta thống nhất là sẽ để file dữ liệu này cùng cấp với file notebook và khi đọc file thì chỉ truyền vào tên của file. Ngoài ra, bạn cũng cần cho cột `Respondent` (id của người làm khảo sát) làm cột index của `survey_df`.

<!-- grade_id: cell-d46e97079096473c, points: 0.25, grade: true, solution: true -->
```python
survey_df = pd.read_csv('survey_results_public-short.csv',index_col='Respondent')
```

```python
# TEST
survey_df.head()
```

### Dữ liệu có bao nhiêu dòng và bao nhiêu cột? (0.25đ)

Kế đến, bạn tính số dòng và số cột của DataFrame `survey_df` và lần lượt lưu vào biến `num_rows` và `num_cols`. 

<!-- grade_id: cell-8338f3610f6c5002, points: 0, grade: false, solution: true -->
```python
num_rows = len(survey_df)
num_cols = len(survey_df.columns)
```

<!-- grade_id: cell-00f00d1bba235517, points: 0.25, grade: true, solution: false -->
```python
# TEST
assert num_rows == 64461
assert num_cols == 28
```

### Mỗi dòng có ý nghĩa gì? Có vấn đề các dòng có ý nghĩa khác nhau không?

Theo file "README_2020.txt" cũng như theo quan sát sơ bộ về dữ liệu, mỗi dòng trong DataFrame `survey_df` cho biết kết quả làm khảo sát của một người. Có vẻ không có vấn đề các dòng có ý nghĩa khác nhau. 

### Dữ liệu có các dòng bị lặp không? (0.25đ)

Kế đến, bạn tính số dòng có index (id của người làm khảo sát) bị lặp và lưu vào biến `num_duplicated_rows`. Trong nhóm các dòng có index giống nhau thì dòng đầu tiên không tính là bị lặp.

<!-- grade_id: cell-aadbfe12be205faa, points: 0, grade: false, solution: true -->
```python
num_duplicated_rows = len(survey_df.index) - len(survey_df.index.drop_duplicates())        
```

<!-- grade_id: cell-7c456d5495d4904d, points: 0.25, grade: true, solution: false -->
```python
# TEST
assert num_duplicated_rows == 0
```

### Mỗi cột có ý nghĩa gì? (0.25đ)

Để xem ý nghĩa của mỗi cột thì:
- Trước tiên, bạn cần đọc file "survey_results_schema-short.csv" vào DataFrame `col_meaning_df`; bạn cũng cần cho cột "Column" làm cột index. 
- Sau đó, bạn chỉ cần hiển thị DataFrame `col_meaning_df` ra để xem (vụ này khó nên ở dưới mình đã làm cho bạn ở cell có dòng "# TEST" 😉). Tuy nhiên, bạn sẽ thấy ở cột "QuestionText": các chuỗi mô tả bị cắt do quá dài. Do đó, trước khi hiển thị DataFrame `col_meaning_df`, bạn cũng cần chỉnh sao đó để các chuỗi mô tả không bị cắt (vụ này bạn tự search Google, gợi ý: bạn sẽ dùng đến câu lệnh `pd.set_option`).   

<!-- grade_id: cell-5efb0517df207535, points: 0.25, grade: true, solution: true -->
```python
col_meaning_df = pd.read_csv('survey_results_schema-short.csv',index_col='Column')
pd.set_option('display.max_colwidth',None)
```

```python
# TEST
col_meaning_df
```

Trước khi đi tiếp, bạn nên đọc kết quả hiển thị ở trên và đảm bảo là bạn đã hiểu ý nghĩa của các cột. Để hiểu ý nghĩa của cột, có thể bạn sẽ cần xem thêm các giá trị của cột bên DataFrame `survey_df`. 

### Mỗi cột hiện đang có kiểu dữ liệu gì? Có cột nào có kiểu dữ liệu chưa phù hợp để có thể xử lý tiếp không? (0.25đ)

Kế đến, bạn tính kiểu dữ liệu (dtype) của mỗi cột trong DataFrame `survey_df` và lưu kết quả vào Series `dtypes` (Series này có index là tên cột). 

<!-- grade_id: cell-206c19691b1a6c05, points: 0, grade: false, solution: true -->
```python
dtypes = survey_df.dtypes
dtypes
```

<!-- grade_id: cell-3ebcf8f07733605b, points: 0.25, grade: true, solution: false -->
```python
# TEST
float_cols = set(dtypes[(dtypes==np.float32) | (dtypes==np.float64)].index)
assert float_cols == {'Age', 'ConvertedComp', 'WorkWeekHrs'}
object_cols = set(dtypes[dtypes == object].index)
assert len(object_cols) == 25
```

Như bạn có thể thấy, cột "YearsCode" và "YearsCodePro" nên có kiểu dữ liệu số, nhưng hiện giờ đang có kiểu dữ liệu object. Ta hãy thử xem thêm về các giá trị 2 cột này. 

```python
survey_df['YearsCode'].unique()
```

```python
survey_df['YearsCodePro'].unique()
```

Ta nên đưa 2 cột này về dạng số để có thể tiếp tục khám phá (tính min, median, max, ...).

---

## Tiền xử lý (0.5đ)

Bạn sẽ thực hiện tiền xử lý để chuyển 2 cột "YearsCode" và "YearsCodePro" về dạng số (float). Trong đó: "Less than 1 year" $\to$ 0, "More than 50 years" $\to$ 51. Sau khi chuyển thì `survey_df.dtypes` sẽ thay đổi.

<!-- grade_id: cell-cfc7d3af0fb5edec, points: 0, grade: false, solution: true -->
```python
# chuyển đổi col YearsCode
survey_df['YearsCode'] = survey_df['YearsCode'].replace('Less than 1 year',0)
survey_df['YearsCode'] = survey_df['YearsCode'].replace('More than 50 years',51)
survey_df['YearsCode'] = pd.to_numeric(survey_df['YearsCode'])
# chuyển đổi col YearsCodePro
survey_df['YearsCodePro'] = survey_df['YearsCodePro'].replace('Less than 1 year',0)
survey_df['YearsCodePro'] = survey_df['YearsCodePro'].replace('More than 50 years',51)
survey_df['YearsCodePro'] = pd.to_numeric(survey_df['YearsCodePro'])

```

<!-- grade_id: cell-4ee41f84604837a1, points: 0.5, grade: true, solution: false -->
```python
# TEST
assert survey_df['YearsCode'].dtype in [np.float32, np.float64]
assert survey_df['YearsCodePro'].dtype in [np.float32, np.float64]
```

---

## Quay lại bước khám phá dữ liệu 

### Với mỗi cột có kiểu dữ liệu dạng số, các giá trị được phân bố như thế nào? (1đ)

(Trong đó: phần tính các mô tả của mỗi cột chiếm 0.5đ, phần tính số lượng giá trị không hợp lệ của mỗi cột chiếm 0.5đ)

Với các cột có kiểu dữ liệu số, bạn sẽ tính:
- Tỉ lệ % (từ 0 đến 100) các giá trị thiếu 
- Giá trị min
- Giá trị lower quartile (phân vị 25)
- Giá trị median (phân vị 50)
- Giá trị upper quartile (phân vị 75)
- Giá trị max

Bạn sẽ lưu kết quả vào DataFrame `nume_col_info_df`, trong đó: 
- Tên của các cột là tên của các cột số trong `survey_df`
- Tên của các dòng là: "missing_percentage", "min", "lower_quartile", "median", "upper_quartile", "max"  

Để dễ nhìn, tất cả các giá trị bạn đều làm tròn với 1 chữ số thập phân bằng phương thức `.round(1)`.

<!-- grade_id: cell-d1503f01712c873a, points: 0, grade: false, solution: true -->
```python
cols_nume_name = list(column for column in survey_df.columns if survey_df[column].dtype.kind in 'iufc' )
index_name = list(["missing_percentage", "min", "lower_quartile", "median", "upper_quartile", "max"])
# tạo dataframe từ cols và index đã cho
nume_col_info_df = pd.DataFrame(0,index=index_name,columns=cols_nume_name)
# xử lý missing_percentage
missing_percentage = (survey_df[cols_nume_name].isna().sum() / len(survey_df) * 100).round(1)
nume_col_info_df.loc['missing_percentage'] = missing_percentage.tolist()
# xử lý min
min = survey_df[cols_nume_name].min()
nume_col_info_df.loc['min'] = min.tolist()
# xử lý lower_quartile
lower_quartile = survey_df[cols_nume_name].quantile(0.25)
nume_col_info_df.loc['lower_quartile'] = lower_quartile.tolist()
# xử lý median
median = survey_df[cols_nume_name].median()
nume_col_info_df.loc['median'] = median.tolist()
# xử lý upper_quartile
upper_quartile = survey_df[cols_nume_name].quantile(0.75)
nume_col_info_df.loc['upper_quartile'] = upper_quartile.tolist()
# xử lý max
max = survey_df[cols_nume_name].max()
nume_col_info_df.loc['max'] = max.tolist()

nume_col_info_df
```

<!-- grade_id: cell-b19aff83344102f8, points: 0.5, grade: true, solution: false -->
```python
# TEST
assert nume_col_info_df.shape == (6, 5)
data = nume_col_info_df.loc[['missing_percentage', 'min', 'lower_quartile', 'median', 'upper_quartile', 'max'],
                           ['Age', 'ConvertedComp', 'WorkWeekHrs', 'YearsCode', 'YearsCodePro']].values
correct_data = np.array([[     29.5,      46.1,      36.2,      10.5,      28.1],
                         [      1. ,       0. ,       1. ,       0. ,       0. ],
                         [     24. ,   24648. ,      40. ,       6. ,       3. ],
                         [     29. ,   54049. ,      40. ,      10. ,       6. ],
                         [     35. ,   95000. ,      44. ,      17. ,      12. ],
                         [    279. , 2000000. ,     475. ,      51. ,      51. ]])
assert np.array_equal(data, correct_data)
```

**Có giá trị không hợp lệ trong mỗi cột không? (không xét giá trị thiếu)**
- Cột "Age": bạn hãy tính số lượng giá trị không hợp lệ của cột "Age" (< giá trị tương ứng trong cột "YearsCode" HOẶC < giá trị tương ứng trong cột "YearsCodePro") và lưu kết quả vào biến `num_invalid_Age_vals`.
- Cột "WorkWeekHrs" (số giờ làm việc trung bình một tuần): ta thấy max là 475 giờ! Trong khi đó, 7 ngày * 24 giờ = 168 giờ! Bạn hãy tính số lượng giá trị không hợp lệ của cột "WorkWeekHrs" (> 24 * 7) và lưu kết quả vào biến `num_invalid_WorkWeekHrs`.
- Cột "YearsCode": bạn hãy tính số lượng giá trị không hợp lệ của cột "YearsCode" (< giá trị tương ứng trong cột "YearsCodePro" HOẶC > giá trị tương ứng trong cột "Age") và lưu kết quả vào biến `num_invalid_YearsCode`.
- Cột "YearsCodePro": bạn hãy tính số lượng giá trị không hợp lệ của cột "YearsCodePro" (> giá trị tương ứng trong cột "YearsCode" HOẶC > giá trị tương ứng trong cột "Age") và lưu kết quả vào biến `num_invalid_YearsCodePro`.

<!-- grade_id: cell-59a58bdc5f54bb25, points: 0, grade: false, solution: true -->
```python
# đếm số row Age không hợp lệ
num_invalid_Age_vals = survey_df.loc[(survey_df['Age'] < survey_df['YearsCode']) | (survey_df['Age'] < survey_df['YearsCodePro'])].Age.count()
# đếm số row WorkWeekHrs không hợp lệ
num_invalid_WorkWeekHrs_vals = survey_df.loc[survey_df['WorkWeekHrs'] > 24*7].WorkWeekHrs.count()
# đếm số row YearsCode không hợp lệ
num_invalid_YearsCode_vals = survey_df.loc[(survey_df['YearsCode'] < survey_df['YearsCodePro']) | (survey_df['YearsCode'] > survey_df['Age'])].YearsCode.count()
# đếm số row YearsCodePro không hợp lệ
num_invalid_YearsCodePro_vals = survey_df.loc[(survey_df['YearsCode'] < survey_df['YearsCodePro']) | (survey_df['YearsCodePro'] > survey_df['Age'])].YearsCodePro.count()
```

<!-- grade_id: cell-386769471c5b4b8f, points: 0.5, grade: true, solution: false -->
```python
# TEST
assert num_invalid_WorkWeekHrs_vals == 62
assert num_invalid_Age_vals == 16
assert num_invalid_YearsCode_vals == 499
assert num_invalid_YearsCodePro_vals == 486
```

Do số lượng các giá trị không hợp lệ cũng khá ít nên ta có thể tiền xử lý bằng cách xóa các dòng chứa các giá trị không hợp lệ.

---

## Tiền xử lý (0.5đ)

Bạn sẽ thực hiện tiền xử lý để xóa đi các dòng của DataFrame `survey_df` mà chứa ít nhất là một giá trị không hợp lệ. Sau khi tiền xử lý thì `survey_df` sẽ thay đổi. 

<!-- grade_id: cell-6a6242ef6d2c8c65, points: 0, grade: false, solution: true -->
```python
# tạo list để lưu vị trí các index không hợp lệ
list_index = []
list_index.extend(survey_df.loc[(survey_df['Age'] < survey_df['YearsCode']) | (survey_df['Age'] < survey_df['YearsCodePro'])].Age.index.tolist())
list_index.extend(survey_df.loc[survey_df['WorkWeekHrs'] > 24*7].WorkWeekHrs.index.tolist())
list_index.extend(survey_df.loc[(survey_df['YearsCode'] < survey_df['YearsCodePro']) | (survey_df['YearsCode'] > survey_df['Age'])].YearsCode.index.tolist())
list_index.extend(survey_df.loc[(survey_df['YearsCode'] < survey_df['YearsCodePro']) | (survey_df['YearsCode'] > survey_df['Age'])].YearsCodePro.index.tolist())
list_index= pd.unique(list_index)
# xáo các rows theo index không hợp lệ
survey_df = survey_df.drop(list_index)

```

<!-- grade_id: cell-9d04f770468db3c8, points: 0.5, grade: true, solution: false -->
```python
# TEST
assert len(survey_df) == 63900
```

---

## Quay lại bước khám phá dữ liệu

### Với mỗi cột có kiểu dữ liệu không phải dạng số, các giá trị được phân bố như thế nào? (1đ)

Với các cột có kiểu dữ liệu không phải số, bạn sẽ tính:
- Tỉ lệ % (từ 0 đến 100) các giá trị thiếu 
- Số lượng các giá trị (các giá trị ở đây là các giá trị khác nhau và ta không xét giá trị thiếu): với cột mà ứng với câu hỏi dạng multichoice (ví dụ, cột "DevType"), mỗi giá trị có thể chứa nhiều choice (các choice được phân tách bởi dấu chấm phẩy), và việc đếm trực tiếp các giá trị không có nhiều ý nghĩa lắm vì số lượng tổ hợp các choice là khá nhiều; một cách khác tốt hơn mà bạn sẽ làm là đếm số lượng các choice
- Tỉ lệ % (từ 0 đến 100) của mỗi giá trị được sort theo tỉ lệ % giảm dần (ta không xét giá trị thiếu, tỉ lệ là tỉ lệ so với số lượng các giá trị không thiếu): bạn dùng dictionary để lưu, key là giá trị, value là tỉ lệ %; với cột mà ứng với câu hỏi dạng multichoice, cách làm tương tự như ở trên

Bạn sẽ lưu kết quả vào DataFrame `cate_col_info_df`, trong đó: 
- Tên của các cột là tên của các cột không phải số trong `survey_df`
- Tên của các dòng là: "missing_percentage", "num_values", "value_percentages"  

Để dễ nhìn, tất cả các giá trị bạn đều làm tròn với 1 chữ số thập phân bằng phương thức `.round(1)`.

Gợi ý: có thể bạn sẽ muốn dùng [phương thức `explode`](https://pandas.pydata.org/pandas-docs/stable/reference/api/pandas.Series.explode.html).

```python
# Các cột ứng với câu hỏi khảo sát multichoice
multichoice_cols = ['DevType', 'Gender', 'JobFactors', 
                    'LanguageWorkedWith', 'LanguageDesireNextYear', 
                    'MiscTechWorkedWith', 'MiscTechDesireNextYear', 
                    'NEWCollabToolsWorkedWith', 'NEWCollabToolsDesireNextYear', 
                    'PlatformWorkedWith', 'PlatformDesireNextYear', 
                    'NEWStuck']
```

<!-- grade_id: cell-c1cfa0a88177079d, points: 0, grade: false, solution: true -->
```python
pd.set_option('display.max_colwidth', 100) # Để dễ nhìn
pd.set_option('display.max_columns', None) # Để dễ nhìn
# tên các cột là numeric
# ở đây ta xét, integer(i), unsigned int(u), float(f), complex(c) là numeric
cols_nume_name = list(column for column in survey_df.columns if survey_df[column].dtype.kind in 'iufc' )
#lưu tên các columns không phải numeric
cols_cate_name = []
for i in survey_df.columns:
    if i not in cols_nume_name:
        cols_cate_name.append(i)
index_name = ["missing_percentage", "num_values", "value_percentages"]
#tạo dataframe từ cols và index đã cho
cate_col_info_df = pd.DataFrame(0,index=index_name,columns=cols_cate_name)

# xử lý missing_percentage
missing_percentage = (survey_df[cols_cate_name].isna().sum() / len(survey_df) * 100).round(1)
cate_col_info_df.loc['missing_percentage'] = missing_percentage.tolist()

#xử lý num_values
#xử lý value_percentages
num_values = []
value_percentages = []
for i in cols_cate_name:
    if i in multichoice_cols:
        temp = survey_df[i].dropna().str.split(';').explode() 
        num_values.append(len(set(temp)))
        temp1 = (temp.value_counts(normalize=True)*100).round(1)
        value_percentages.append(temp1.to_dict())
    else:
        temp = survey_df[i].dropna().explode()
        num_values.append(len(set(temp)))
        temp1 = (temp.value_counts(normalize=True)*100).round(1)
        value_percentages.append(temp1.to_dict())
cate_col_info_df.loc['num_values'] = num_values
cate_col_info_df.loc['value_percentages'] = value_percentages

cate_col_info_df
```

<!-- grade_id: cell-e748038de60d2ab7, points: 1, grade: true, solution: false -->
```python
# TEST
c = cate_col_info_df['MainBranch']
assert c.loc['missing_percentage'] == 0.5
assert c.loc['num_values'] == 5
assert c.loc['value_percentages']['I am a developer by profession'] == 73.5

c = cate_col_info_df['Hobbyist']
assert c.loc['missing_percentage'] == 0.1
assert c.loc['num_values'] == 2
assert c.loc['value_percentages']['Yes'] == 78.2

c = cate_col_info_df['DevType']
assert c.loc['missing_percentage'] == 23.6
assert c.loc['num_values'] == 23
assert c.loc['value_percentages']['Academic researcher'] == 2.2

c = cate_col_info_df['PlatformWorkedWith']
assert c.loc['missing_percentage'] == 16.5
assert c.loc['num_values'] == 16
assert c.loc['value_percentages']['Docker'] == 10.6
```

---

## Đặt câu hỏi 

Sau khi khám phá dữ liệu, ta đã hiểu hơn về dữ liệu. Bây giờ, ta hãy xem thử có câu hỏi nào có thể được trả lời bằng dữ liệu này.

**Một câu hỏi có thể có là:** Platform nào (Windows, Linux, Docker, AWS, ...) được yêu thích nhất, platform nào được yêu thích nhì, platform nào được yêu thích ba, ...?

Một platform được xem là được yêu thích nếu một người đã dùng platform này (cột "PlatformWorkedWith") và muốn tiếp tục dùng platform trong năm kế (cột "PlatformDesireNextYear").

**Trả lời được câu hỏi này sẽ** phần nào giúp ta định hướng là nên tập trung học platform nào để có thể chuẩn bị cho tương lai (mình nói "phần nào" vì ở đây dữ liệu chỉ giới hạn trong phạm vi những người làm khảo sát của StackOverflow).

---

## Tiền xử lý

Nếu bạn thấy cần thực hiện thêm thao tác tiền xử lý để chuẩn bị dữ liệu cho bước phân tích thì bạn làm ở đây. Bước này là không bắt buộc. 

```python
#explode 2 column thành 2 series độc lập, bỏ hết các giá trị na
s1 = survey_df['PlatformWorkedWith'].dropna().str.split(';').explode()
s2 = survey_df['PlatformDesireNextYear'].dropna().str.split(';').explode()
```

---

## Phân tích dữ liệu (2.25đ)

Bây giờ, bạn sẽ thực hiện phân tích dữ liệu để trả lời cho câu hỏi ở trên. Cụ thể các bước như sau:
- Bước 1: tính Series `most_loved_platforms`, trong đó: 
    - Index là tên flatform (ở bước khám phá dữ liệu, bạn đã thấy có tất cả 16 platform)
    - Data là tỉ lệ % (từ 0 đến 100, được làm tròn với một chữ số thập phân bằng phương thức `round(1)`) được yêu thích (được sort giảm dần) 
- Bước 2: từ Series `most_loved_platforms`, bạn vẽ bar chart:
    - Bạn cho các bar nằm ngang (cho dễ nhìn)
    - Bạn đặt tên trục hoành là "Tỉ lệ %"

Code bước 1.

<!-- grade_id: cell-6d34c88e2fd46b33, points: 0, grade: false, solution: true -->
```python
# Chuyển 2 series về dataframe, giữ nguyên index là index của 2 series (append = True)
df1 = s1.to_frame().set_index('PlatformWorkedWith', append=True)
df2 = s2.to_frame().set_index('PlatformDesireNextYear', append=True)
# intersect 2 dataframe theo index thành MultiIndex
idx = df1.index.intersection(df2.index)
# chuyển MultiIndex về thành series, sau đó làm tròn
most_loved_platforms = (pd.Series(idx.get_level_values(1), idx.get_level_values(0)).value_counts(normalize=True)*100).round(1)
most_loved_platforms
```

<!-- grade_id: cell-1d22896d0f3f2394, points: 2, grade: true, solution: false -->
```python
# TEST
assert len(most_loved_platforms) == 16
assert most_loved_platforms.loc['Linux'] == 20.2
assert most_loved_platforms.loc['Windows'] == 14.6
assert most_loved_platforms.loc['Docker'] == 12.3
```

Code bước 2.

<!-- grade_id: cell-5395c611d6065339, points: 0.25, grade: true, solution: true -->
```python
most_loved_platforms.plot.barh(x='Ti le %',figsize=(20,10))
plt.xlabel('Ti le %')
plt.ylabel('Platform')
```

Bạn đã hiểu tại sao mình khuyên bạn là nên tập làm quen dần với các câu lệnh của Linux chưa 😉 

---

## Đặt câu hỏi của bạn (1.5đ)

Bây giờ, đến lượt bạn phải tự suy nghĩ và đưa ra câu hỏi mà có thể trả lời bằng dữ liệu. Ngoài việc đưa ra câu hỏi, bạn cũng phải giải thích để người đọc thấy nếu trả lời được câu hỏi thì sẽ có lợi ích gì. Bạn nên sáng tạo một xíu, không nên đưa ra câu hỏi cùng dạng với câu hỏi của mình ở trên.  

<!-- grade_id: cell-2e7a92ae99ca7449, points: 1.5, grade: true, solution: true -->
**Câu hỏi:** Thu nhập,số lượng người dùng dựa theo những ngôn ngữ lập trình mà người khảo sát đã từng học: Ngôn ngữ lập trình nào có thể cho thu nhập trung bình tốt nhất, tốt nhì,...Ngôn ngữ nào được nhiều người dùng nhất, nhì,...
- Ở đây nếu người được khảo sát đã học nhiều ngôn ngữ lập trình tính đến thời điểm khảo sát, thì ta coi như từng ngôn ngữ đó đều có thu nhập bằng thu nhập của người khảo sát.

**Trả lời được câu hỏi này** có thể giúp chúng ta hiểu hơn hiện tại ngôn ngữ lập trình nào có thể có thu nhập trung bình tốt nhất, ngôn ngữ nào được sử dụng phổ biến nhất, từ đó giúp ta định hướng được có nên học thêm ngôn ngữ khác trong tương lai không.


---

## Tiền xử lý để chuẩn bị dữ liệu cho bước phân tích để trả lời cho câu hỏi của bạn

Phần này là không bắt buộc. 

```python
# bỏ hết các giá trị na()
df_temp = survey_df[["LanguageWorkedWith","ConvertedComp"]].dropna()
# chuyển các chuỗi ở column LanguageWorkedWith thành list
df_temp["LanguageWorkedWith"] = df_temp["LanguageWorkedWith"].str.split(";")
# explode để có được datafram với columns LanguageWorkedWith chỉ chưa 1 ngôn ngữ duy nhất
df_temp = df_temp.explode("LanguageWorkedWith")
```

---

## Phân tích dữ liệu để ra câu trả lời cho câu hỏi của bạn (2đ)

### Phân tích dữ liệu

- Bước 1: Tính Dataframe `df_language` trong đó:
    - index là tên của các language có được sau khi explode column `LanguageWorkedWith` của `survey_df` 
    - 2 columns của Dataframe là mean(trung bình) và count(số lượng)
- Bước 2: từ Dataframe  `df_language` vẽ scatter plot với x là `average annual income ($)` và y là `number user of language`

- Bước 1:

<!-- grade_id: cell-83e11892184cf811, points: 2, grade: true, solution: true -->
```python
# groupby theo tên ngôn ngữ sau đó tính mean, làm tròn
temp1 = df_temp.groupby(["LanguageWorkedWith"]).mean().round(1)
# đặt lại tên cột cho Dataframe temp1
temp1.columns = ['mean']
# groupby theo tên ngôn ngữ sau đó tính số lượng
temp2 = df_temp.groupby(["LanguageWorkedWith"]).count()
# đặt lại tên cột cho Dataframe temp1
temp2.columns = ['count']
# ghép 2 dataframe vừa tạo được ở trên thành 1 dataframe dựa vào các index trùng nhau
df_result = pd.concat([temp1,temp2],axis=1)
```

- Bước 2:

```python
# Tạo ngẫu nhiên list các màu của các chấm trên biểu đồ
color_dots = ["#"+''.join([np.random.choice(list('0123456789ABCDEF')) for j in range(6)])
             for i in range(len(df_result))]
# Chỉnh size cho biểu đồ
plt.figure(figsize = (20,10),dpi = 80)
# Vẽ biêt đồ scatter
plot = plt.scatter(df_result['mean'],df_result['count'],c=color_dots)
# Đặt tên cho các dot
for i,txt in enumerate(df_result.index):
    plt.annotate(txt,(df_result['mean'][i],df_result['count'][i]))
# Đặt title, tên trục x, tên trục y
plt.title('number user and average anual income of programing language')
plt.xlabel('average  annual income ($/year)')
plt.ylabel('number user of language')

```

***Nhận xét:***  Có thể thấy Julia có vẻ đang là ngôn ngữ cho được thu nhập trung bình tốt nhất nhưng lại có số người dùng rất ít do đang còn là ngôn ngữ mới và hiện nay chỉ thấy dùng trong các ngành nghiên cứu khoa học. JavaScript là ngôn ngữ được dùng nhiều nhất, cũng như nhóm ngôn ngữ chuyên về thiết kế website nhưng mức lương thì có vẻ nằm dưới trung bình.

Những ngôn ngữ thiên về thiết kế web thì hiện tại đang có số người sử dụng lớn, nhưng cùng với việc được nhiều người sử dụng nên thu nhập của những ngôn ngữ này không quá cao, còn những ngôn ngữ mới phục vụ cho việc nghiên cứu khoa học, AI tuy có số lượng người dùng ít nhưng lại có thu nhập trung bình cao. 

Từ đó nếu muốn học theo xu hướng ngôn ngữ hiện tại, cộng đồng đã phát triển lớn mạnh thì nên học website như javaScript, HTML/CSS, SQL, còn nếu muốn thu nhập cao thì nên đi theo hướng nghiên cứu khoa học và học những ngôn ngữ mới phục vụ tốt cho nó như julia, ruby, scala, peri