# HW2: Numpy

*Vì đây là bài tập về Numpy nên yêu cầu nói chung là không được dùng vòng lặp; còn ở chỗ nào mà cho phép dùng thì mình sẽ nói rõ*

(Cập nhật lần cuối: 28/11/2021)



Họ tên: Lưu Trường Dương

MSSV: 19120489

---

## Cách làm bài và nộp bài

&#9889; Bạn lưu ý là mình sẽ dùng chương trình hỗ trợ chấm bài nên bạn cần phải tuân thủ chính xác qui định mà mình đặt ra, nếu không rõ thì hỏi, chứ không nên tự tiện làm theo ý của cá nhân.

**Cách làm bài**

Bạn sẽ làm trực tiếp trên file notebook này. Đầu tiên, bạn điền họ tên và MSSV vào phần đầu file ở bên trên. Trong file, bạn làm bài ở những chỗ có ghi là:
```python
# YOUR CODE HERE
raise NotImplementedError()
```
hoặc đối với những phần code không bắt buộc thì là:
```python
# YOUR CODE HERE (OPTION)
```
hoặc đối với markdown cell thì là:
```markdown
YOUR ANSWER HERE
```
Tất nhiên, khi làm thì bạn xóa dòng `raise NotImplementedError()` đi.
Đối những phần yêu cầu code thì thường ở ngay phía dưới sẽ có một (hoặc một số) cell chứa các bộ test để giúp bạn biết đã code đúng hay chưa; nếu chạy cell này không có lỗi gì thì có nghĩa là qua được các bộ test. Trong một số trường hợp, các bộ test có thể sẽ không đầy đủ; nghĩa là, nếu không qua được test thì là code sai, nhưng nếu qua được test thì chưa chắc đã đúng hoàn toàn.

Trong khi làm bài, bạn có thể cho in ra màn hình, tạo thêm các cell để test. Nhưng khi nộp bài thì bạn xóa các cell mà bạn tự tạo, xóa hoặc comment các câu lệnh in ra màn hình. Bạn lưu ý <font color=red>không được tự tiện xóa các cell hay sửa code của Thầy</font> (trừ những chỗ được phép sửa như đã nói ở trên).

Bạn cố gắng giữ cho code rõ ràng, ngăn nắp: đặt tên biến gợi ý, comment ở những chỗ cần comment, giữ cho một dòng code không quá dài (vì nếu không thì người đọc sẽ phải mất công scroll, bạn có thể dùng qui tắc: một dòng code chỉ nên có tối đa là 79 ký tự). 

Trong khi làm bài, thường xuyên `Ctrl + S` để lưu lại bài làm của bạn, tránh mất mát thông tin.


*Nên nhớ mục tiêu chính ở đây là <font color=green>học, học một cách chân thật</font>. Bạn có thể thảo luận ý tưởng với bạn khác, nhưng <font color=green>code và bài làm phải là của bạn, dựa trên sự hiểu thật sự của bạn</font>. <font color=red>Nếu vi phạm thì sẽ bị 0 điểm cho toàn bộ môn học.</font>*

**Cách nộp bài**

Khi chấm bài, đầu tiên mình sẽ chọn `Kernel` - `Restart & Run All` để restart và chạy tất cả các cell trong notebook của bạn; do đó, trước khi nộp bài, bạn nên chạy thử `Kernel` - `Restart & Run All` để đảm bảo mọi chuyện diễn ra đúng như mong đợi.

Sau đó, bạn tạo thư mục nộp bài theo cấu trúc sau:
- Thư mục `MSSV` (vd, nếu bạn có MSSV là 1234567 thì bạn đặt tên thư mục là `1234567`)
    - File `HW2.ipynb` (không cần nộp các file khác)

Cuối cùng, bạn nén thư mục `MSSV` này lại và nộp ở link trên moodle. <font color=red>Bạn lưu ý tuân thủ chính xác qui định nộp bài.</font>

---

## Môi trường code

Ta thống nhất trong môn này:
- Code trên môi trường Linux (bạn đã setup ở HW0)
- Dùng Jupyter Notebook hoặc Jupyter Lab để code file .ipynb. Lưu ý: bạn <font color=red>không dùng Google Colab</font>, vì Google Colab có thể xóa một số thông tin meta mà chương trình hỗ trợ chấm bài của mình đã tạo trong file bài tập
- Dùng phiên bản các package như trong file "min_ds-env.yml" (bạn xem trong HW0 cách dùng Conda để tạo môi trường code từ file .yml và cách kích hoạt môi trường này sau khi đã tạo)

```python
import sys
sys.executable
```

Output:
```
'E:\\app\\python\\python.exe'
```

Nếu không có vấn đề gì thì file chạy python sẽ là file của môi trường code "min_ds-env".

Trong bài này, bạn không dùng thư viện Pandas.

---

## Import các thư viện cần thiết

```python
import numpy as np
# YOUR CODE HERE (OPTION)
```

---

## Thu thập dữ liệu

Trong bài này, ta chỉ đơn giản là dùng bộ dữ liệu đã được thu thập sẵn. Bạn download [ở đây](https://files.grouplens.org/datasets/movielens/ml-100k/u.data); bạn để tên file là "u.data", không có đuôi txt hay đuôi gì cả. Bộ dữ liệu này thật ra có nhiều hơn một file, nhưng trong bài này ta chỉ dùng một file là "u.data".

Theo [file mô tả dữ liệu](https://files.grouplens.org/datasets/movielens/ml-100k-README.txt):
> MovieLens data sets were collected by the GroupLens Research Project
at the University of Minnesota.
> 
> This data set consists of:
> * 100,000 ratings (1-5) from 943 users on 1682 movies. 
> * Each user has rated at least 20 movies.  
> * Simple demographic info for the users (age, gender, occupation, zip)
>
> The data was collected through the MovieLens web site
(movielens.umn.edu) during the seven-month period from September 19th, 
1997 through April 22nd, 1998. This data has been cleaned up - users
who had less than 20 ratings or did not have complete demographic
information were removed from this data set. Detailed descriptions of
the data file can be found at the end of this file.

Trên là mô tả chung cho tất cả các file trong bộ dữ liệu. Còn đây là mô tả cụ thể cho file "u.data":
> The full u data set, 100000 ratings by 943 users on 1682 items. Each user has rated at least 20 movies. Users and items are numbered consecutively from 1. The data is randomly ordered. This is a tab separated list of user id | item id | rating | timestamp. The time stamps are unix seconds since 1/1/1970 UTC 

---

## Khám phá dữ liệu

### Dữ liệu có bao nhiêu dòng và bao nhiêu cột? (0.5đ)

Việc đầu tiên của bạn là đọc file "u.data" vào Numpy array tên là `raw_ratings` (dùng hàm `np.genfromtxt`). Bạn để `dtype` là `np.uint64`. Bạn cũng cần để file "u.data" cùng thư mục với file notebook này và khi gọi hàm `np.genfromtxt` thì chỉ truyền vào tên file là "u.data".

Kế tiếp, bạn tính số dòng và số cột, rồi lần lượt lưu vào hai biến là `n_rows` và `n_cols`.

<!-- grade_id: cell-1bf7d4be2f1eb278, points: 0, grade: false, solution: true -->
```python
raw_ratings = np.genfromtxt("u.data",dtype = np.uint64)
n_rows = raw_ratings.shape[0]
n_cols = raw_ratings.shape[1]
```

<!-- grade_id: cell-fd69b593a2722fbc, points: 0.5, grade: true, solution: false -->
```python
# TEST
assert raw_ratings.dtype == np.uint64
assert raw_ratings.ndim == 2
assert n_rows == 100000
assert n_cols == 4
raw_ratings[:5] # Xem 5 dòng đầu tiên
```

Output:
```
array([[      196,       242,         3, 881250949],
       [      186,       302,         3, 891717742],
       [       22,       377,         1, 878887116],
       [      244,        51,         2, 880606923],
       [      166,       346,         1, 886397596]], dtype=uint64)
```

### Mỗi dòng có ý nghĩa gì? Có vấn đề các dòng có ý nghĩa khác nhau không?

Theo file mô tả dữ liệu và theo quan sát sơ bộ về dữ liệu thì một dòng cho biết thông tin cho điểm của một người dùng đối với một phim. Có vẻ không có dòng nào bị "lạc loài".

### Dữ liệu có các dòng bị lặp không? (0.5đ)

Bạn sẽ kiểm tra vụ này và lưu kết quả vào biến `have_duplicated_rows`. Biến này sẽ có giá trị True nếu dữ liệu có các dòng bị lặp và có giá trị False nếu ngược.

<!-- grade_id: cell-c184a3a04425a1c1, points: 0, grade: false, solution: true -->
```python
have_duplicated_rows = False if  len(np.unique(raw_ratings,axis=0)) == len(raw_ratings) else True
```

<!-- grade_id: cell-7f2d2f011cc33f7d, points: 0.5, grade: true, solution: false -->
```python
# TEST
assert have_duplicated_rows == False
```

### Mỗi cột có ý nghĩa gì?

Theo file mô tả dữ liệu thì:
- Cột thứ nhất cho biết id của người dùng
- Cột thứ hai cho biết id của phim
- Cột thứ ba cho biết điểm mà người dùng đã cho đối với phim
- Cột thứ tư cho biết thời điểm mà người dùng cho điểm (được thể hiện bằng số giây tính từ một cột mốc chuẩn)

### Mỗi cột hiện đang có kiểu dữ liệu gì? Có cột nào có kiểu dữ liệu chưa phù hợp để có thể xử lý tiếp không?

Hiện tại thì tất cả các cột đều có kiểu dữ liệu là số:

```python
raw_ratings.dtype
```

Output:
```
dtype('uint64')
```

Tất cả các cột đều có kiểu dữ liệu phù hợp. 

Bạn đang nghĩ cột thời điểm cho điểm có kiểu dữ liệu phù hợp hay chưa? Cột này đang thể hiện cho thời gian dưới dạng số (số giây tính từ một cột mốc chuẩn). Ta vẫn có thể tính min, max, ... bình thường với cột này; và khi cần thì ta có thể chuyển min, max, ... sang dạng ngày tháng năm để dễ nhìn. Tất nhiên, ngay lúc này, ta có thể chuyển cột này sang dạng ngày tháng năm luôn; rồi khi tính min, max, ... thì kết quả sẽ ở sẵn dạng ngày tháng năm. Tuy nhiên, ở đây, mình không chọn cách làm này vì ràng buộc của Numpy array là các cột phải có cùng kiểu dữ liệu; nếu làm theo cách này thì dữ liệu của sẽ bị phân làm 2 array: một ứng với cột thời điểm cho điểm, một ứng với các cột còn lại (như vậy thì sẽ hơi bất tiện).

### Với mỗi cột có kiểu dữ liệu dạng số (numeric) và dạng datetime, các giá trị được phân bố như thế nào? (1đ)

Nhìn qua thì có vẻ tất cả các cột đều có dạng numeric. Nhưng mình nghĩ nên phân cột "user id" và "movie id" (2 cột đầu) vào nhóm categorical. Theo cách làm này thì sẽ có 2 cột dạng numeric: cột điểm (cột chỉ số 2) và cột thời điểm cho điểm (cột chỉ số 3).

2 cột numeric này có bao nhiêu giá trị thiếu?

```python
np.sum(np.isnan(raw_ratings[:, 2:]), axis=0)
```

Output:
```
array([0, 0])
```

Việc của bạn là tính min, median, max của 2 cột numeric này (dùng hàm `np.percentile`, để tham số `interpolation` mặc định là `linear`); rồi lần lượt lưu vào 2 biến là `rate_col_profile` và `rate_date_col_profile`. Đây là 2 Numpy array (một chiều), trong đó `rate_col_profile` có `dtype` là số thực, còn `rate_date_col_profile` có `dtype` là `'datetime64[s]'`.

Gợi ý: khi ép kiểu một array thì bạn có thể truyền vào kiểu cần ép dưới dạng một string (ở ngay trên cũng có một string kìa ;-) ).

<!-- grade_id: cell-15364ea0e0e9c792, points: 0, grade: false, solution: true -->
```python
rate_col_profile = np.array([np.min(raw_ratings[:,2]),np.percentile(raw_ratings[:,2],50),\
    np.max(raw_ratings[:,2])],dtype='f')
rate_date_col_profile = np.array([np.min(raw_ratings[:,3]),np.percentile(raw_ratings[:,3],50),\
    np.max(raw_ratings[:,3])],dtype='datetime64[s]')
```

<!-- grade_id: cell-3f3b0054dd6eadb9, points: 1, grade: true, solution: false -->
```python
# TEST
assert np.array_equal(rate_col_profile, np.array([1., 4., 5.]))
assert np.array_equal(rate_date_col_profile, 
                      np.array(['1997-09-20T03:05:10', 
                                '1997-12-22T21:42:24',
                                '1998-04-22T23:10:38'], 
                               dtype='datetime64[s]'))
```

Min và max của 2 cột này không có gì bất thường, khớp với file mô tả dữ liệu.

### Với mỗi cột có kiểu dữ liệu dạng phân loại (categorical), các giá trị được phân bố như thế nào? (1đ)

Như đã nói, 2 cột categorical là 2 cột đầu.

2 cột này có bao nhiêu giá trị thiếu:

```python
np.sum(np.isnan(raw_ratings[:, :2]), axis=0)
```

Output:
```
array([0, 0])
```

Việc của bạn là, với mỗi cột, tính ra list gồm 5 con số: số lượng giá trị khác nhau, giá trị xuất hiện ít nhất cùng số lượng xuất hiện (tổng cộng 2 con số), giá trị xuất hiện nhiều nhất cùng số lượng xuất hiện (tổng cộng 2 con số). Bạn lưu 2 list tính được ứng với 2 cột vào 2 biến lần lượt là `user_col_profile` và `movie_col_profile`.

Ta thống nhất: nếu có nhiều giá trị cùng xuất hiện ít nhất thì lấy giá trị nhỏ nhất, nếu có nhiều giá trị cùng xuất hiện nhiều nhất thì lấy giá trị lớn nhất (gợi ý này có thể giúp ích cho bạn: trong các hàm sort của Numpy, nếu bạn cho giá trị của tham số `kind` là `stable` thì Numpy sẽ sort array với đảm bảo các giá trị giống nhau trong array ban đầu sẽ được giữ nguyên thứ tự tương đối với nhau trong array kết quả). 

<!-- grade_id: cell-09819763bd956c3a, points: 0, grade: false, solution: true -->
```python
user_col_profile = []
movie_col_profile = []


#xử lý user_col_profile 
temp = raw_ratings[:,0]
user_col_profile.append(len(np.unique(temp)))

temp1,temp2 = np.unique(temp,return_counts=True)

indexs_min = np.where(temp2 == np.min(temp2))

user_col_profile.append(temp1[indexs_min[0][0]])
user_col_profile.append(np.min(temp2))

indexs_max = np.where(temp2 == np.max(temp2))

user_col_profile.append(temp1[indexs_max[0][-1]])
user_col_profile.append(np.max(temp2))

# #xử lý movie_col_profile
temp = raw_ratings[:,1]
movie_col_profile.append(len(np.unique(temp)))

temp1,temp2 = np.unique(temp,return_counts=True)

indexs_min = np.where(temp2 == np.min(temp2))

movie_col_profile.append(temp1[indexs_min[0][0]])
movie_col_profile.append(np.min(temp2))

indexs_max = np.where(temp2 == np.max(temp2))

movie_col_profile.append(temp1[indexs_max[0][-1]])
movie_col_profile.append(np.max(temp2))

```

<!-- grade_id: cell-9f4bbdb1bcb99063, points: 1, grade: true, solution: false -->
```python
assert user_col_profile==[943, # Có chừng này user
                          19,  # Đây là user rate ít movie nhất
                          20,  # và đó là chừng này movie
                          405, # Đây là user rate nhiều movie nhất
                          737] # và đó là chừng này movie
assert movie_col_profile==[1682,#Có chừng này movie
                           599, #Đây là movie được ít user rate nhất
                           1,   #và đó là chừng này user
                           50,  #Đây là movie được nhiều user rate nhất
                           583] #và đó là chừng này user
```

Các thông tin này khớp với file mô tả dữ liệu:
> The full u data set, 100000 ratings by 943 users on 1682 items. Each user has rated at least 20 movies. Users and items are numbered consecutively from 1.

Ý "Users and items are numbered consecutively from 1" chưa được kiểm chứng từ kết quả ở trên. Dưới đây, ta sẽ kiểm chứng ý này:

```python
print('User id  - min & max:', 
      raw_ratings[:, 0].min(), '&', raw_ratings[:, 0].max()) 
print('Movie id - min & max:', 
      raw_ratings[:, 1].min(), '&', raw_ratings[:, 1].max())
```

Output:
```
User id  - min & max: 1 & 943
Movie id - min & max: 1 & 1682
```

Như vậy là hoàn toàn khớp (đúng không?).

---

## Đưa ra câu hỏi có ý nghĩa cần trả lời

Đến đây, ta đã hiểu hơn về dữ liệu. Bây giờ, ta sẽ xem thử có câu hỏi nào có ý nghĩa mà có thể được trả lời bằng dữ liệu này không.

Một câu hỏi có thể có là: \
*Với mỗi người dùng, ta nên đề xuất (recommend) các phim nào để người dùng xem?*

Việc tìm ra câu trả lời cho câu hỏi này sẽ giúp ích cho cả người dùng và nhà cung cấp dịch vụ xem phim:

- Người dùng: người dùng có thể muốn xem phim, nhưng có rất nhiều phim và người dùng không biết nên xem phim nào; sẽ rất tiện cho người dùng nếu hệ thống có thể đề xuất một danh sách phim mà nhiều khả năng người dùng sẽ thích
- Nhà cung cấp dịch vụ xem phim: nếu hệ thống đề xuất tốt thì nhiều khả năng người dùng sẽ xem phim và sẽ thích, nghĩa là người dùng sẽ tiếp tục trả phí để sử dụng dịch vụ ;-)

---

## Tiền xử lý (1đ)

Việc xây dựng hệ thống đề xuất sản phẩm mình đã có nói sơ bộ ở buổi học 09. Đầu tiên, từ dữ liệu ban đầu là array `raw_ratings`, bạn sẽ cần tạo array `ratings` gồm số-lượng-user dòng và số-lượng-movie cột, `ratings[i, j]` cho biết user có id i đã cho movie có id j bao nhiêu điểm, nếu chưa cho thì giá trị sẽ là `nan`. Vì Python đánh chỉ số từ 0 nên bạn cũng cần chỉnh lại để id của user và movie bắt đầu từ 0 (thay vì 1).

<!-- grade_id: cell-11ff284c3fe7fb7e, points: 0, grade: false, solution: true -->
```python
# tạo ma trận chứa toàn các giá trị nan
ratings = np.empty((len(set(raw_ratings[:,0])),len(set(raw_ratings[:,1]))))
ratings[:] = np.nan

temp = raw_ratings.copy()

temp[:,0] -=1
temp[:,1] -=1

ratings[temp[:,0],temp[:,1]] = temp[:,2]
```

<!-- grade_id: cell-fb18066caaa407ed, points: 1, grade: true, solution: false -->
```python
assert ratings.shape == (943, 1682)
missing_ratios = np.mean(np.isnan(ratings))
assert missing_ratios.round(3) == 0.937 # Có 93.7% giá trị nan!
```

---

## Phân tích dữ liệu để trả lời cho câu hỏi (6đ)

Để có thể đề xuất các movie cho mỗi user thì ta sẽ phải làm 2 công đoạn:

- Công đoạn 1: dự đoán điểm mà mỗi user sẽ cho đối với các movie mà user chưa xem (nói một cách khác: điền các giá trị nan trong array `ratings`)
- Công đoạn 2: đề xuất cho user các movie (mà user chưa xem) có điểm dự đoán cao

Trong bài này thì ta chỉ tập trung làm công đoạn 1.

Ở buổi học 09, mình đã hướng dẫn cách dự đoán điểm cho một user. Trong bài này, bạn sẽ làm với tất cả user. Tuy nhiên, bạn sẽ không dùng vòng lặp và lần lượt làm cho từng user vì Numpy không thích vòng lặp ;-). Thay vào đó, bạn sẽ làm một lúc cho tất cả user. Để làm được điều này thì với mỗi user bạn cần tính điểm (rating) cho tất cả các movie; nếu bạn chỉ tính điểm cho các movie mà chưa được cho điểm như ở buổi học 09 thì sẽ khó có thể làm một lúc cho tất cả user được, vì mỗi user sẽ có số lượng movie chưa cho điểm khác nhau. Việc tính điểm cho các movie đã có điểm là thừa (điểm được tính lại thật ra sẽ rất gần với điểm ban đầu do điểm ban đầu sẽ có trọng số lớn nhất), nhưng sẽ có khả năng giúp tăng tốc vì ta có thể làm với tất cả user cùng một lúc.

Nếu thử làm thì bạn sẽ thấy vẫn không thể làm với tất cả user cùng một lúc được, do các array kết quả trung gian sẽ có kích thước khoảng số-lượng-user x số-lượng-user x số-lượng-movie, và sẽ bị báo không đủ bộ nhớ. Một cách để vượt qua vấn đề này mà không quay về cách duyệt từng user là duyệt một lúc một nhóm các user (gọi là một batch) và dự đoán điểm của nhóm các user này; kích thước của các array kết quả trung gian sẽ giảm xuống khoảng số-lượng-user-của-một-batch x số-lượng-user x số-lượng-movie. Như vậy, ở chỗ này bạn sẽ được phép dùng vòng lặp để duyệt từng batch (nhóm user).

Bạn sẽ dùng `batch_size` (số lượng user của một nhóm) là 32. Và sau cùng, bạn sẽ điền cho array `filled_ratings` điểm của tất cả movie; trong array này, mình chỉ quan tâm đến các giá trị điền ứng với các giá trị nan trong mảng `ratings` ban đầu (còn các giá trị khác bạn cho sao cũng được). Trong quá trình làm, bạn lưu ý không làm thay đổi array `ratings`.

```python
batch_size = 32
filled_ratings = np.empty_like(ratings)
```

Mình không có ý làm bạn hoảng sợ nhưng việc cài đặt sẽ khá phức tạp, đỏi hỏi nắm vững cơ chế broadcasting và khả năng tưởng tượng 3 chiều cao ;-). Do đó, đầu tiên bạn sẽ thử làm với một batch trước. Mình sẽ hỗ trợ bạn bằng cách chia ra thành các bước và cung cấp bộ test ở mỗi bước. Ý quan trọng mà bạn cần nắm khi chuyển từ làm với một user (như ở buổi học 09) sang làm với `batch_size` user là: các array sẽ có thêm một chiều ở đầu với kích thước là `batch_size` ứng với `batch_size` user. \
Sau khi code đã chạy đúng với một batch, bạn sẽ làm thật với tất cả dữ liệu.

Đầu tiên, bạn sẽ thử làm với một batch ứng với các user có chỉ số từ `start` đến `end` (5.5đ).

```python
start = 0
end = batch_size
```

Bước 1: tính array `similarities` cho biết độ tương tự giữa mỗi user trong batch đang xét với tất cả các user trong toàn bộ dữ liệu (cách tính độ tương tự giống như ở buổi học 09). Array này sẽ có kích thước `batch_size` x `n_users` (`n_users` là số lượng user trong toàn bộ dữ liệu), `similarities[i, j]` sẽ cho biết độ tương tự giữa user chỉ số i với user chỉ số j. Trong trường hợp 2 user không có cho điểm chung movie nào (khi chạy bạn có thể sẽ thấy warning "RuntimeWarning: Mean of empty slice") thì bạn cho độ tương tự bằng 0.

<!-- grade_id: cell-1d34779ee4335381, points: 0, grade: false, solution: true -->
```python
temp = np.abs(ratings - ratings[start:end, None,:])
temp1 = np.nanmean(temp, axis=2)
similarities = 1/(0.001 + temp1)
# chuyển nan về 0
similarities = np.nan_to_num(similarities)
```

Output:
```
C:\Users\luutr\AppData\Local\Temp/ipykernel_10016/3873511762.py:2: RuntimeWarning: Mean of empty slice
  temp1 = np.nanmean(temp, axis=2)
```

<!-- grade_id: cell-3c3f855070577541, points: 2, grade: true, solution: false -->
```python
# TEST
assert similarities.shape == (32, 943)
assert np.array_equal(similarities[:3, :3].round(1), 
                      np.array([[1.0e+03, 1.1e+00, 7.0e-01],
                                [1.1e+00, 1.0e+03, 7.0e-01],
                                [7.0e-01, 7.0e-01, 1.0e+03]]))
```

```python
assert np.array_equal(similarities[:5, :5].round(1),
np.array([[1.0e+03, 1.1e+00, 7.0e-01, 8.0e-01, 1.0e+00],
[1.1e+00, 1.0e+03, 7.0e-01, 7.0e-01, 1.7e+00],
[7.0e-01, 7.0e-01, 1.0e+03, 5.0e-01, 1.0e+00],
[8.0e-01, 7.0e-01, 5.0e-01, 1.0e+03, 2.0e+00],
[1.0e+00, 1.7e+00, 1.0e+00, 2.0e+00, 1.0e+03]]))
```

Bước 2: tính array trọng số `weights` là array độ tương tự `similarity` đã được chuẩn hóa. Với mỗi user trong batch đang xét, độ tương tự với tất cả các user sẽ được chuẩn hóa theo từng movie (giống như ở buổi học 09): với các user mà đã cho điểm movie thì chuẩn hóa bằng cách chia độ tương tự cho tổng độ tương tự mà chỉ xét trên các user đã cho điểm movie, với các user còn lại chuẩn hóa bằng cách cho độ tương tự bằng 0. Array `weights` sẽ có kích thước là `batch_size` x `n_users` x `n_movies` (với `n_movies` là tổng số lượng movie). \
Khi chạy, bạn sẽ thấy warning "RuntimeWarning: invalid value encountered in true_divide"; đây là do các user mà có cho điểm một movie đang xét đều có độ tương tự bằng 0 với một user đang xét, dẫn đến khi chuẩn hóa sẽ ra 0/0 và kết quả là nan. Trường hợp này có nghĩa là không đủ thông tin để dự đoán điểm và trong bài này thì bạn cứ tạm để như vậy.

```python
w = similarities.copy()
normalized_w = w[:,:,None] * ~np.isnan(ratings)
weights = normalized_w/ np.sum(normalized_w,axis=1)[:,None,:]
weights
```

Output:
```
C:\Users\luutr\AppData\Local\Temp/ipykernel_10016/2588024967.py:3: RuntimeWarning: invalid value encountered in true_divide
  weights = normalized_w/ np.sum(normalized_w,axis=1)[:,None,:]
array([[[6.84033520e-01, 8.80920564e-01, 9.17640778e-01, ...,
         0.00000000e+00, 0.00000000e+00, 0.00000000e+00],
        [7.23504722e-04, 0.00000000e+00, 0.00000000e+00, ...,
         0.00000000e+00, 0.00000000e+00, 0.00000000e+00],
        [0.00000000e+00, 0.00000000e+00, 0.00000000e+00, ...,
         0.00000000e+00, 0.00000000e+00, 0.00000000e+00],
        ...,
        [9.75796747e-04, 0.00000000e+00, 0.00000000e+00, ...,
         0.00000000e+00, 0.00000000e+00, 0.00000000e+00],
        [0.00000000e+00, 0.00000000e+00, 0.00000000e+00, ...,
         0.00000000e+00, 0.00000000e+00, 0.00000000e+00],
        [0.00000000e+00, 7.68974546e-04, 0.00000000e+00, ...,
         0.00000000e+00, 0.00000000e+00, 0.00000000e+00]],

       [[6.69930549e-04, 6.59369283e-03, 9.68620263e-03, ...,
         0.0000
```

<!-- grade_id: cell-04b8a47e95e8dd09, points: 2, grade: true, solution: false -->
```python
# TEST
assert weights.shape == (32, 943, 1682)
assert np.sum(np.isnan(weights)) == 31119
```

```python
assert weights[0].sum().round() == 1682
assert weights[1].sum().round() == 1682
```

Bước 3: Với mỗi user trong batch đang xét, tính điểm (cho tất cả movie) bằng cách nhân điểm của tất cả user với trọng số tương ứng trong array `weight`; rồi ghi các điểm của mỗi user xuống một dòng trong array `filled_ratings`. 

<!-- grade_id: cell-e73c0a38ffa925db, points: 0, grade: false, solution: true -->
```python
filled_ratings[start:end] = np.nansum(ratings * weights, axis=1)
```

<!-- grade_id: cell-9a094032ccb1726a, points: 1.5, grade: true, solution: false -->
```python
# TEST
filled_batch = filled_ratings[start:end]
filled_nanvals = filled_batch[np.isnan(ratings[start:end])]
assert np.array_equal(filled_nanvals[:3].round(1),
                      np.array([3.6, 3.5, 4. ]))
assert np.array_equal(filled_nanvals[-3:].round(1),
                      np.array([2., 3., 3. ]))
```

Code của bạn đã chạy OK với một batch rồi thì bây giờ bạn có thể làm thật trên toàn bộ dữ liệu (0.5đ).

<!-- grade_id: cell-9ee6c3bd8e2dea97, points: 0, grade: false, solution: true -->
```python
filled_ratings = np.empty_like(ratings)

batch_size =32
start = 0
end = batch_size

flag = True

if end > len(filled_ratings):
    end = len(filled_ratings)

while flag:
    temp = np.abs(ratings - ratings[start:end, None,:])
    temp1 = np.nanmean(temp, axis=2)
    similarities = np.reciprocal(0.001 + temp1)
    similarities = np.nan_to_num(similarities)

    normalized_w = similarities[:,:,None] * ~np.isnan(ratings)
    weights = normalized_w/ np.sum(normalized_w,axis=1)[:,None,:]

    filled_ratings[start:end] = np.nansum(ratings * weights, axis=1)

    start=end
    end +=batch_size
    if start == len(filled_ratings):
        flag = False
    if end > len(filled_ratings):
        end = len(filled_ratings)
        
```

Output:
```
C:\Users\luutr\AppData\Local\Temp/ipykernel_9340/1404763321.py:14: RuntimeWarning: Mean of empty slice
  temp1 = np.nanmean(temp, axis=2)
C:\Users\luutr\AppData\Local\Temp/ipykernel_9340/1404763321.py:19: RuntimeWarning: invalid value encountered in true_divide
  weights = normalized_w/ np.sum(normalized_w,axis=1)[:,None,:]
```

<!-- grade_id: cell-5e6aaf23d47e5193, points: 0.5, grade: true, solution: false -->
```python
# TEST
filled_nanvals = filled_ratings[np.isnan(ratings)]
assert np.array_equal(filled_nanvals[:3].round(1),
                      np.array([3.6, 3.5, 4. ]))
assert np.array_equal(filled_nanvals[-3:].round(1),
                      np.array([0., 3., 3. ]))
```